### Make a list of the classic 'ROGYBIV' rainbow colours (1 point)

<!-- grade_id: cell-5b443cfa56e38e51, points: 0, grade: false, solution: true -->
```python
colours = []

##BEGIN SOLUTION
colours = ["red", "orange", "yellow", "green", "blue", "indigo", "violet"]
#END SOLUTION
```

<!-- grade_id: cell-fa56848035077eea, points: 1, grade: true, solution: false -->
```python
assert colours[0] == "red"
assert colours[1] == "orange"
assert colours[2] == "yellow"
assert colours[3] == "green"
assert colours[4] == "blue"
assert colours[5] == "indigo"
assert colours[6] == "violet"

```